rule_files:
  - /prometheus/rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'up{container="cos-fleetshard-operator-camel", namespace="some_namespace"}'
        values: '1+0x30 0+0x30 1+0x30'
      - series: 'kube_pod_container_status_restarts_total{container="cos-fleetshard-operator-camel", namespace="some_namespace"}'
        values: '0+0x30 1+1x30 0+0x100'

    alert_rule_test:
      # CosFleetShardOperatorCamelTargetDown test units
      - eval_time: 10m
        alertname: CosFleetShardOperatorCamelTargetDown
        exp_alerts: [ ]
      - eval_time: 50m
        alertname: CosFleetShardOperatorCamelTargetDown
        exp_alerts:
          - exp_labels:
              alertname: CosFleetShardOperatorCamelTargetDown
              severity: critical
              container: 'cos-fleetshard-operator-camel'
              namespace: 'some_namespace'
            exp_annotations:
              summary: 'the cos-fleetshard-operator-camel target is down'
              description: 'the cos-fleetshard-operator-camel target has been unable to scrape the cos-fleetshard-operator-camel container in the some_namespace namespace for longer than 10 minutes'
              sop_url: 'https://github.com/bf2fc6cc711aee1a0c2a/cos-tools/blob/main/observability/sops/alerts/container_down.adoc'
      - eval_time: 70m
        alertname: CosFleetShardOperatorCamelTargetDown
        exp_alerts: [ ]

      # CosFleetShardOperatorCamelContainerFrequentlyRestarting test units
      - eval_time: 10m
        alertname: CosFleetShardOperatorCamelContainerFrequentlyRestarting
        exp_alerts: [ ]
      - eval_time: 50m
        alertname: CosFleetShardOperatorCamelContainerFrequentlyRestarting
        exp_alerts:
          - exp_labels:
              alertname: CosFleetShardOperatorCamelContainerFrequentlyRestarting
              severity: critical
              container: 'cos-fleetshard-operator-camel'
              namespace: 'some_namespace'
            exp_annotations:
              summary: 'the cos-fleetshard-operator-camel operator is restarting frequently'
              description: 'the cos-fleetshard-operator-camel operator container restarted frequently in the last 60 minutes'
              sop_url: 'https://github.com/bf2fc6cc711aee1a0c2a/cos-tools/blob/main/observability/sops/alerts/container_down.adoc'
      - eval_time: 160m
        alertname: CosFleetShardOperatorCamelContainerFrequentlyRestarting
        exp_alerts: [ ]