rule_files:
  - /prometheus/rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'up{container="cos-fleetshard-operator-debezium", namespace="some_namespace"}'
        values: '1+0x30 0+0x30 1+0x30'
      - series: 'kube_pod_container_status_restarts_total{container="cos-fleetshard-operator-debezium", namespace="some_namespace"}'
        values: '0+0x30 1+1x30 0+0x100'

    alert_rule_test:
      # CosFleetShardOperatorDebeziumTargetDown test units
      - eval_time: 10m
        alertname: CosFleetShardOperatorDebeziumTargetDown
        exp_alerts: [ ]
      - eval_time: 50m
        alertname: CosFleetShardOperatorDebeziumTargetDown
        exp_alerts:
          - exp_labels:
              alertname: CosFleetShardOperatorDebeziumTargetDown
              severity: critical
              container: 'cos-fleetshard-operator-debezium'
              namespace: 'some_namespace'
            exp_annotations:
              summary: 'the cos-fleetshard-operator-debezium target is down'
              description: 'the cos-fleetshard-operator-debezium target has been unable to scrape the cos-fleetshard-operator-debezium container in the some_namespace namespace for longer than 10 minutes'
              sop_url: 'https://github.com/bf2fc6cc711aee1a0c2a/cos-tools/blob/main/observability/sops/alerts/cos_fleet_shard_operator_debezium_target_down.adoc'
      - eval_time: 70m
        alertname: CosFleetShardOperatorDebeziumTargetDown
        exp_alerts: [ ]

      # CosFleetShardOperatordebeziumContainerFrequentlyRestarting test units
      - eval_time: 10m
        alertname: CosFleetShardOperatorDebeziumContainerFrequentlyRestarting
        exp_alerts: [ ]
      - eval_time: 50m
        alertname: CosFleetShardOperatorDebeziumContainerFrequentlyRestarting
        exp_alerts:
          - exp_labels:
              alertname: CosFleetShardOperatorDebeziumContainerFrequentlyRestarting
              severity: critical
              container: 'cos-fleetshard-operator-debezium'
              namespace: 'some_namespace'
            exp_annotations:
              summary: 'the cos-fleetshard-operator-debezium operator is restarting frequently'
              description: 'the cos-fleetshard-operator-debezium operator container restarted frequently in the last 60 minutes'
              sop_url: 'https://github.com/bf2fc6cc711aee1a0c2a/cos-tools/blob/main/observability/sops/alerts/cos_fleet_shard_operator_debezium_container_frequently_restarting.adoc'
      - eval_time: 160m
        alertname: CosFleetShardOperatorDebeziumContainerFrequentlyRestarting
        exp_alerts: [ ]